[1613561396] SERVICE ALERT: localhost;HTTP;WARNING;SOFT;1;HTTP WARNING: HTTP/1.1 401 Unauthorized - 691 bytes in 0.001 second response time
[1613561456] SERVICE ALERT: localhost;HTTP;WARNING;SOFT;2;HTTP WARNING: HTTP/1.1 401 Unauthorized - 691 bytes in 0.001 second response time
[1613561509] SERVICE ALERT: localhost;SSH;CRITICAL;SOFT;1;connect to address 127.0.0.1 and port 22: Connection refused
[1613561516] SERVICE ALERT: localhost;HTTP;WARNING;SOFT;3;HTTP WARNING: HTTP/1.1 401 Unauthorized - 691 bytes in 0.001 second response time
[1613561546] SERVICE ALERT: localhost;Swap Usage;CRITICAL;SOFT;1;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1613561569] SERVICE ALERT: localhost;SSH;CRITICAL;SOFT;2;connect to address 127.0.0.1 and port 22: Connection refused
[1613561576] SERVICE ALERT: localhost;HTTP;WARNING;HARD;4;HTTP WARNING: HTTP/1.1 401 Unauthorized - 691 bytes in 0.002 second response time
[1613561606] SERVICE ALERT: localhost;Swap Usage;CRITICAL;SOFT;2;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1613561629] SERVICE ALERT: localhost;SSH;CRITICAL;SOFT;3;connect to address 127.0.0.1 and port 22: Connection refused
[1613561666] SERVICE ALERT: localhost;Swap Usage;CRITICAL;SOFT;3;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1613561689] SERVICE ALERT: localhost;SSH;CRITICAL;HARD;4;connect to address 127.0.0.1 and port 22: Connection refused
[1613561726] SERVICE ALERT: localhost;Swap Usage;CRITICAL;HARD;4;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1613567336] SERVICE ALERT: viewnext113;HTTP;CRITICAL;SOFT;1;connect to address 10.250.13.1 and port 80: Connection refused
[1613567396] SERVICE ALERT: viewnext113;HTTP;CRITICAL;SOFT;2;connect to address 10.250.13.1 and port 80: Connection refused
[1613567486] SERVICE ALERT: viewnext113;Swap Usage;CRITICAL;SOFT;1;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1613567516] SERVICE ALERT: viewnext113;HTTP;CRITICAL;SOFT;3;connect to address 10.250.13.1 and port 80: Connection refused
[1613567546] SERVICE ALERT: viewnext113;Swap Usage;CRITICAL;SOFT;2;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1613567576] SERVICE ALERT: viewnext113;HTTP;CRITICAL;HARD;4;connect to address 10.250.13.1 and port 80: Connection refused
[1613567606] SERVICE ALERT: viewnext113;Swap Usage;CRITICAL;SOFT;3;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1613567666] SERVICE ALERT: viewnext113;Swap Usage;CRITICAL;HARD;4;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1613572373] SERVICE ALERT: viewadmin110;HTTP;CRITICAL;SOFT;1;connect to address 10.250.13.1 and port 80: Connection refused
[1613572433] SERVICE ALERT: viewadmin110;HTTP;CRITICAL;SOFT;2;connect to address 10.250.13.1 and port 80: Connection refused
[1613572493] SERVICE ALERT: viewadmin110;HTTP;CRITICAL;SOFT;3;connect to address 10.250.13.1 and port 80: Connection refused
[1613572533] SERVICE ALERT: viewadmin110;Swap Usage;CRITICAL;SOFT;1;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1613572553] SERVICE ALERT: viewadmin110;HTTP;CRITICAL;HARD;4;connect to address 10.250.13.1 and port 80: Connection refused
[1613572593] SERVICE ALERT: viewadmin110;Swap Usage;CRITICAL;SOFT;2;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1613572653] SERVICE ALERT: viewadmin110;Swap Usage;CRITICAL;SOFT;3;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1613572713] SERVICE ALERT: viewadmin110;Swap Usage;CRITICAL;HARD;4;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1613574035] SERVICE ALERT: viewnext113;HTTP;CRITICAL;SOFT;1;connect to address 10.250.13.1 and port 80: Connection refused
[1613574104] SERVICE ALERT: viewnext113;Jenkins;WARNING;SOFT;1;HTTP WARNING: HTTP/1.1 401 Unauthorized - 782 bytes in 0.155 second response time
[1613574155] SERVICE ALERT: viewnext113;HTTP;CRITICAL;SOFT;2;connect to address 10.250.13.1 and port 80: Connection refused
[1613574224] SERVICE ALERT: viewnext113;Jenkins;WARNING;SOFT;2;HTTP WARNING: HTTP/1.1 401 Unauthorized - 782 bytes in 0.171 second response time
[1613574275] SERVICE ALERT: viewnext113;HTTP;CRITICAL;HARD;3;connect to address 10.250.13.1 and port 80: Connection refused
[1613574344] SERVICE ALERT: viewnext113;Jenkins;WARNING;HARD;3;HTTP WARNING: HTTP/1.1 401 Unauthorized - 782 bytes in 0.242 second response time
[1613574954] SERVICE ALERT: viewnext113;Jenkins;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1613575554] SERVICE ALERT: viewnext113;Jenkins;OK;HARD;3;HTTP OK: HTTP/1.1 200 OK - 52236 bytes in 0.271 second response time